## Instructions :
- Only Edit  *YOUR CODE HERE*
- Do not change the assert statements which test your functions

```python
import numpy as np
```

###Q. Windows and Strides
Given a 1d array, and parameters window_len and stride_len, do the following:
- Divide the 1d array into smaller arrays of size window_len starting at index 0
- move the starting index by step size stride_len each time
- Return all windows as a 2d matrix

Eg.,
Input: <br>
arr = [1, 3, 7, 1, 2, 6, 0, 1]  <br>
stride_len = 2 , window_len = 4

- First window (starting at index 0) should be of len window_len (4) --> [1, 3, 7, 1] <br>
- Move by step size stride_len(2) and take next window --> [7, 1, 2, 6] <br>
and so on..

Output matrix <br>
[[1, 3, 7, 1], <br>
 [7, 1, 2, 6], <br>
 [2, 6, 0, 1]]


<!-- grade_id: cell-01751c442d3d9ac2, points: 0, grade: false, solution: true -->
```python
import numpy as np
def gen_strides(a, stride_len, window_len):
    '''
    Input:
      a: Numpy array of 1 dimension
      stride_len: int, stride length
      window_len : int, window length
    
    Output:
      Numpy array of 2 dimensions containing windowed strides as explained above
    '''
    # YOUR CODE HERE
    nrows = ((a.size-window_len)//stride_len)+1 
    new = np.array([]) 
    for i in range (0,len(a), stride_len):
        if (len(a[i:])>=window_len):
            new = np.append(new, a[i:i+window_len]) 
    new = new.reshape((nrows, window_len)) 
    return (new)
```

<!-- grade_id: cell-096efc2d6c969658, points: 2, grade: true, solution: false -->
```python
"""Test for strides"""
assert(np.all(gen_strides(np.array([1, 3, 7, 1, 2, 6, 0, 1]),2,4) == [[1, 3, 7, 1], [7, 1, 2, 6], [2, 6, 0, 1]]))

print("Sample Tests passed", '\U0001F44D')
```

Output:
```
Sample Tests passed 👍
```

###Q. Shuffle
Given a numpy array of arbitrary dimensions (> 1), shuffle its rows randomly. <br>
Hint - You need to shuffle along axis 0

<!-- grade_id: cell-9d0cd8b9d15e98b5, points: 0, grade: true, solution: true -->
```python
import random
def shuf(arr):
  '''
  Input: 
    arr: Numpy array of arbitrary number of dimensions (>1)
  Output:
    numpy array of same shape as arr but with rows shuffled
  '''
  # YOUR CODE HERE
  np.random.shuffle(arr)
  return arr

```

<!-- grade_id: cell-9c47b636484fca56, points: 2, grade: true, solution: false -->
```python
"""Test for shuf"""
assert np.any(shuf(np.array([[1, 2, 3],[4, 5, 6],[7, 8, 9]])) != np.array([[1, 2, 3], [4, 5, 6], [7, 8, 9]]))
assert shuf(np.array([[1, 2, 3],[4, 5, 6],[7, 8, 9]])).shape == np.array([[1, 2, 3], [4, 5, 6], [7, 8, 9]]).shape

print("Sample Tests passed", '\U0001F44D')
```

Output:
```
Sample Tests passed 👍
```

###Q. Match
    Get the positions where corresponding elements (same indices) of array a and array b match <br>
    
    Eg. Input <br>
    a = np.array([1,2,3,2,3,4,3,4,5,6]) <br>
    b = np.array([7,2,10,2,7,4,9,4,9,8]) <br>
    
    Desired Output:

    [1, 3, 5, 7])

<!-- grade_id: cell-b0d486da07272b4e, points: 0, grade: false, solution: true -->
```python
def match(a,b):
    '''
    Inputs:
      a, b: numpy arrays of same shape of 1 dimension
    Outputs:
      list containing indices where both arrays have same elements
    '''
    # YOUR CODE HERE
    match_list = []
    for i in range(a.shape[0]):
      if(a[i] == b[i]):
        match_list.append(i)
    return match_list


```

<!-- grade_id: cell-cfe91a664ba7f193, points: 2, grade: true, solution: false -->
```python
"""Test for match"""
assert(match(np.array([1,2,3,2,3,4,3,4,5,6]),np.array([7,2,10,2,7,4,9,4,9,8])) == [1,3,5,7])
print("Sample Tests passed", '\U0001F44D')
```

Output:
```
Sample Tests passed 👍
```

#Q. Inverse of Array

<!-- grade_id: cell-454bd2de95987a3f, points: 0, grade: true, solution: true -->
```python
def inv(arr):
    """
    Given an array arr (square matrix), find its inverse
    """
    # YOUR CODE HERE
    return np.linalg.inv(arr)



```

<!-- grade_id: cell-43abef740a3b78fb, points: 2, grade: true, solution: false -->
```python
"""Test for inv"""
assert np.all(np.isclose(inv(np.array([[6, 1, 1], [4, -2, 5], [2, 8, 7]])).tolist(), np.array([[0.17647058823529413, -0.0032679738562091526, -0.02287581699346405],[0.05882352941176469, -0.130718954248366, 0.0849673202614379],[-0.1176470588235294, 0.1503267973856209, 0.0522875816993464]])))

print("Sample Tests passed", '\U0001F44D')

```

Output:
```
Sample Tests passed 👍
```

#Q. System of linear equations

<!-- grade_id: cell-202ca6eaf85ae4c3, points: 0, grade: false, solution: true -->
```python
def lin_eqn(a,b):
    '''
    Solve the system of linear equations
    of the form ax = b
    
    Eg. 
    
    Solve the system of linear equation
    
    x + 2*y = 8
    3*x + 4*y = 18
    
    Given inputs a and b represent coefficients and constant of linear equation respectively
    
    coefficients: 
    a = np.array([[1, 2], [3, 4]]) 
    
    constants: 
    b = np.array([8, 18])
    
    Desired Output: [2,3]
    

    '''
    # YOUR CODE HERE
    return np.linalg.solve(a,b)
```

<!-- grade_id: cell-648f81e951895cdb, points: 2, grade: true, solution: false -->
```python
"""Test for lin_eqn"""
assert np.all(np.isclose(lin_eqn(np.array([[1, 2], [3, 4]]),np.array([8, 18])).tolist(), np.array([2.0000000000000004, 2.9999999999999996]), atol = 1e-5))

print("Sample Tests passed", '\U0001F44D')

```

Output:
```
Sample Tests passed 👍
```

###Q. Rank Array
Rank the items in a multidimensional array 'arr'. <br>
The rank of an item is its index in the sorted list of all items in 'arr' (starting from 0). <br>
          Eg., consider arr = [0, 6, 14, 12, 11] <br>
                  arr_sorted = [0, 6, 11, 12, 14] <br>
                  So rank of 0 -> 0; 6->1; 11->2, 12->3, 14->4 <br>
                  rankArray returns a list where each element is replaced by its rank <br>
                  rankArray(arr) -> [0, 1, 4, 3, 2] <br>
                  
==========================================    \\
Another example, <br>
Eg: 

arr = [[ 9  4 15  0 17], [16 17  8  9  0]] ... (2X5 array) <br>

Desired output: <br>
[[4 2 6 0 8], [7 9 3 5 1]] ... (2X5 array) <br>


Here minimum value in arr is 0 at (0,3), so rank the index of 0 in output matrix as 0 <br>
i.e. out[0][3] = 0 <br>
Next minimum is also 0 at index (1,4), so rank the index (1,4) in output matrix as 1 and so on...

============================================= \\
Assume that all elements will be distinct.

<!-- grade_id: cell-642dc0e15679ce42, points: 0, grade: false, solution: true -->
```python
import numpy as np

def rankArray(arr):
    '''
    Input:
        arr: Numpy array of arbitrary dimensions 
    Output:
        numpy array of same shape as arr but with elements replaced by their ranks
    '''
    # YOUR CODE HERE
    rank_arr=arr.ravel().argsort().argsort().reshape(arr.shape).tolist()            
    
    return rank_arr
```

<!-- grade_id: cell-ab95f609e0073276, points: 2, grade: true, solution: false -->
```python
"""Test for rankArray"""
assert(np.all(rankArray(np.array([[9, 4, 15, 0, 17], [16,17,8,9,0]])) == np.array([[4,2, 6, 0, 8], [7, 9, 3, 5, 1]]).tolist()))
print("Sample Tests passed", '\U0001F44D')
```

Output:
```
Sample Tests passed 👍
```